---
title       : SMT201
description : Бие даалтын бодлого 16
attachments :
  slides_link : https://s3.amazonaws.com/smt201-16/san3.pdf

--- type:NormalExercise lang:r xp:100 skills:1 key:0ed836297d

## 16 Бодлого

![equation](https://s3.amazonaws.com/smt201-16/render16_1.png),   n=10  биномт тархалтын **p** -параметрыг 
туршилтын момент $ M_1 = \overline{x} $ = 6.8  -ээр моментын арга хэрэглэн ол.

*** =instructions

- Биномт тархалтын хувьд **I** эрэмбийн анхны момент 
- ![equation](https://s3.amazonaws.com/smt201-16/render16-2.png) тул
- $ 10 \cdot p = 6.8 \Rightarrow p=0.68 $ гэж параметр олдоно.

*** =hint
# Биномт тархалтын хувьд дараах томъёог ашигла.
- ![equation](https://s3.amazonaws.com/smt201-16/render16-2.png)

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Биномт тархалтын хувьд I эрэмбийн төвийн момент
library(polynom)
# ерөнхий нийлбэр
M1 <- 0

# M_1 бодож байна. Олон гишүүнтийн аргумент `x` хувьсагч байна.
for(k in 1:10){ M1 <- M1 + k*choose(10,k)*poly.calc(rep(0,k))*(-1)**k*poly.calc(rep(1,10-k)) }
M1

# **p** параметрын утгыг олвол,
# 10 p = 6.8
p <- 6.8/10

```

*** =solution
```{r}
# Биномт тархалтын хувьд I эрэмбийн төвийн момент
library(polynom)
# ерөнхий нийлбэр
M1 <- 0

# M_1 бодож байна. Олон гишүүнтийн аргумент `x` хувьсагч байна.
for(k in 1:10){ M1 <- M1 + k*choose(10,k)*poly.calc(rep(0,k))*(-1)**k*poly.calc(rep(1,10-k)) }
M1

# **p** параметрын утгыг олвол,
# 10 p = 6.8
p <- 6.8/10

```

*** =sct
```{r}
#no SCT
```
--- type:NormalExercise lang:r xp:100 skills:1 key:922de32376

## 17 Бодлого

Түүврийг ашиглан  **y = a1 x + a2**  с.х -ний параметрууд **a1,a2** ийг хамгийн бага квадратын аргаар тодорхойл.

Хамгийн бага  квадратын аргаар a1, a2 коеффициентуудыг олохдоо

- $ a1 = \frac{M(xy)-MxMy}{Mx^2-(Mx)^2}$
- $ a2 = n \frac{Mx^2My - Mx M(xy)}{Mx^2-(Mx)^2}$
- $ y = a1 \cdot x + a2$ гэж олдоно.

*** =instructions

x, y -ын дундаж:

- `Mx <- mean(x)`
- `My <- mean(y)`

xy, $x^2$ -ын дундаж:

- `M(xy) <- mean(x*y)`
- `M(xx) <- mean(x**2)`

коеффициентуудыг олвол:

- $a1 = \frac{M(xy)-MxMy}{Mx^2-(Mx)^2}$
- $a2 = n \frac{Mx^2My - Mx M(xy)}{Mx^2-(Mx)^2}$
- $ y = a1x + a2$ гэж олдоно.

*** =hint
# no hint

*** =pre_exercise_code
```{r}
x <- c(-1, -0.75, -0.5, -0.25, 0.0, 0.25, 0.5, 0.75, 1.0)
y <- c(2.08, 1.83, 1.57, 1.13, 0.89, 0.75, 0.30, 0.06, -0.01)
n <- 9
```

*** =sample_code
*** =sample_code
```{r}
# x, y` -ын дундаж
Mx <- mean(x)
My <- mean(y)
# xy, $x^2$ -ын дундаж
Mxy <- mean(x*y)
Mxx <- mean(x**2)

# коеффициентуудыг олвол,
a1 <- (Mxy - Mx*My)/(Mxx - Mx**2)
a2 <- n*(Mx**2*My - Mx*Mxy)/(Mxx -Mx**2)
y_lin  <- a1*x + a2  # гэж олдоно.
plot(x,y,type="l",col=2)
lines(x,y_lin,col=3)
```


*** =solution

```{r}
# x, y` -ын дундаж
Mx <- mean(x)
My <- mean(y)
# xy, $x^2$ -ын дундаж
Mxy <- mean(x*y)
Mxx <- mean(x**2)

# коеффициентуудыг олвол,
a1 <- (Mxy - Mx*My)/(Mxx - Mx**2)
a2 <- n*(Mx**2*My - Mx*Mxy)/(Mxx -Mx**2)
y_lin  <- a1*x + a2 # гэж олдоно.
plot(x,y,type="l",col=2)
lines(x,y_lin,col=3)
```

*** =sct
```{r}
#no SCT
```
