---
title       : SMT201
description : Бие даалтын бодлого 16
attachments :
  slides_link : https://s3.amazonaws.com/smt201-16/san3.pdf

--- type:NormalExercise lang:r xp:100 skills:1 key:0ed836297d

## 16 Бодлого

![equation](https://s3.amazonaws.com/smt201-16/render16_1.png),   n=10  биномт тархалтын **p** -параметрыг 
туршилтын момент $ M_1 = \overline{x} $ = 6.8  -ээр моментын арга хэрэглэн ол.

*** =instructions

- Биномт тархалтын хувьд I эрэмбийн төвийн момент 
- ![equation](https://s3.amazonaws.com/smt201-16/render16-2.png) тул
- $ 10 \cdot p = 6.8 \Rightarrow p=0.68 $ гэж параметр олдоно.

*** =hint
# Биномт тархалтын хувьд дараах томъёог ашигла.
- ![equation](https://s3.amazonaws.com/smt201-16/render16-2.png)

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Биномт тархалтын хувьд I эрэмбийн төвийн момент


# **p** параметрын утгыг олвол,
# 10 p = 6.8
p <- 6.8/10

```

*** =sct
```{r}
#no SCT
```
