---
title       : SMT201 IV хэсэг
description : Бие даалтын бодлого 16-20
attachments :
  slides_link : https://s3.amazonaws.com/smt201-16/san3.pdf

--- type:NormalExercise lang:r xp:150 skills:1 key:0ed836297d

## Бодлого 16

![equation](https://s3.amazonaws.com/smt201-16/render16_1.png),   n=10  биномт тархалтын **p** -параметрыг 
туршилтын момент $ M_1 = \overline{x} $ = 6.8  -ээр моментын арга хэрэглэн ол.

*** =instructions

- Биномт тархалтын хувьд **I** эрэмбийн анхны момент 
- ![equation](https://s3.amazonaws.com/smt201-16/render16-2.png) тул
- $ 10 \cdot p = 6.8 \Rightarrow p=0.68 $ гэж параметр олдоно.

*** =hint
# Биномт тархалтын хувьд дараах томъёог ашигла.
- ![equation](https://s3.amazonaws.com/smt201-16/render16-2.png)

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Олон гишүүнтийн санг дуудах
library(polynom)

# Биномт тархалтын хувьд **I** эрэмбийн төвийн момент
M1 <- 0

# M_1 бодож байна. Олон гишүүнтийн аргумент `x` хувьсагч байна.
for(k in 1:10){ M1 <- M1 + k*choose(10,k)*poly.calc(rep(0,k))*(-1)**k*poly.calc(rep(1,10-k)) }
# Үр дүнг харах
M1

# **p** параметрын утгыг олвол,
# 10 p = 6.8
p <- 6.8/10

```

*** =solution
```{r}
# Олон гишүүнтийн санг дуудах
library(polynom)

# Биномт тархалтын хувьд **I** эрэмбийн төвийн момент
M1 <- 0

# M_1 бодож байна. Олон гишүүнтийн аргумент `x` хувьсагч байна.
for(k in 1:10){ M1 <- M1 + k*choose(10,k)*poly.calc(rep(0,k))*(-1)**k*poly.calc(rep(1,10-k)) }
# Үр дүнг харах
M1

# **p** параметрын утгыг олвол,
# 10 p = 6.8
p <- 6.8/10
```

*** =sct
#no SCT


--- type:NormalExercise lang:r xp:150 skills:1 key:922de32376
## Бодлого 17

Түүврийг ашиглан  **y = a1 x + a2**   с.х -ний параметрууд **a1, a2** ийг хамгийн бага квадратын аргаар тодорхойл.

Хамгийн бага  квадратын аргаар a1, a2 коеффициентуудыг олохдоо

- $ a1 = \frac{M(xy)-MxMy}{Mx^2-(Mx)^2}$
- $ a2 = n \frac{Mx^2My - Mx M(xy)}{Mx^2-(Mx)^2}$
- $ y = a1 \cdot x + a2$ гэж олдоно.

*** =instructions
$x, y$ -ын дундаж:

- `Mx <- mean(x)`
- `My <- mean(y)`

$xy, x^2$ -ын дундаж:

- `M(xy) <- mean(x*y)`
- `M(xx) <- mean(x**2)`

Коеффициентуудыг олвол:

- $a1 = \frac{M(xy)-MxMy}{Mx^2-(Mx)^2}$
- $a2 = n \frac{Mx^2My - Mx M(xy)}{Mx^2-(Mx)^2}$
- $ y = a1x + a2$ гэж олдоно.

*** =hint
# no hint

*** =pre_exercise_code
```{r}
x <- c(-1, -0.75, -0.5, -0.25, 0.0, 0.25, 0.5, 0.75, 1.0)
y <- c(2.08, 1.83, 1.57, 1.13, 0.89, 0.75, 0.30, 0.06, -0.01)
n <- 9
```

*** =sample_code
```{r}
# x, y -ын дундаж
Mx <- mean(x)
My <- mean(y)
# xy, x*x -ын дундаж
Mxy <- mean(x*y)
Mxx <- mean(x**2)

# коеффициентуудыг олвол,
a1 <- (Mxy - Mx*My)/(Mxx - Mx**2)
a2 <- n*(Mx**2*My - Mx*Mxy)/(Mxx - Mx**2)

# регрессын тэгшитгэлийг бодож зурвал,
y_lin  <- a1*x + a2
plot(x,y,type="l",col=2)
lines(x,y_lin,col=3)
```


*** =solution

```{r}
# x, y -ын дундаж
Mx <- mean(x)
My <- mean(y)
# xy, x*x -ын дундаж
Mxy <- mean(x*y)
Mxx <- mean(x**2)

# коеффициентуудыг олвол,
a1 <- (Mxy - Mx*My)/(Mxx - Mx**2)
a2 <- n*(Mx**2*My - Mx*Mxy)/(Mxx - Mx**2)

# регрессын тэгшитгэлийг бодож зурвал,
y_lin  <- a1*x + a2
plot(x,y,type="l",col=2)
lines(x,y_lin,col=3)
```

*** =sct
#no SCT

--- type:NormalExercise lang:r xp:150 skills:1 key:c984f046d9
## Бодлого 18

$x1 ,\, x_2, \, x3$ гурвалжны өнцөгүүдээр авсан санамсаргүй хэмжигдэхүүнүүд. 3 өнцгийн үл хамаарсан хэмжилтүүд $\tilde{x}1=17^\circ,\, \tilde{x}2=81^\circ,\, \tilde{x}3=83^\circ$  байсан. Гурвалжны дотоод өнцгийн хувьд $x1 + x2 + x3 = 180^\circ$. Тэгвэл хэмжилтийн алдаануудыг $ N[0, \sigma(k)], \, k =1,2,3$ нормаль тархалттай гэж үзээд хэмжилтийн утгуудыг сайжруул. Энд $ \sigma(i)^2=(0.1\circ)^2, \quad i=1,2,3 $ гэж өгөдөв.

*** =instructions

$x, y$ -ын дундаж:

- `Mx <- mean(x)`
- `My <- mean(y)`

$xy, x^2$ -ын дундаж:

- `M(xy) <- mean(x*y)`
- `M(xx) <- mean(x**2)`

Коеффициентуудыг олвол:

- $a1 = \frac{M(xy)-MxMy}{Mx^2-(Mx)^2}$
- $a2 = n \frac{Mx^2My - Mx M(xy)}{Mx^2-(Mx)^2}$
- $ y = a1x + a2$ гэж олдоно.

*** =hint
# no hint

*** =pre_exercise_code
```{r}
x <- c(-1, -0.75, -0.5, -0.25, 0.0, 0.25, 0.5, 0.75, 1.0)
y <- c(2.08, 1.83, 1.57, 1.13, 0.89, 0.75, 0.30, 0.06, -0.01)
n <- 9
```

*** =sample_code
```{r}
# x, y -ын дундаж
Mx <- mean(x)
My <- mean(y)
# xy, x^2 -ын дундаж
Mxy <- mean(x*y)
Mxx <- mean(x**2)

# коеффициентуудыг олвол,
a1 <- (Mxy - Mx*My)/(Mxx - Mx**2)
a2 <- n*(Mx**2*My - Mx*Mxy)/(Mxx - Mx**2)

# регрессын тэгшитгэлийг бодож зурвал,
y_lin  <- a1*x + a2
plot(x,y,type="l",col=2)
lines(x,y_lin,col=3)
```


*** =solution

```{r}
# x, y -ын дундаж
Mx <- mean(x)
My <- mean(y)
# xy, x*x -ын дундаж
Mxy <- mean(x*y)
Mxx <- mean(x**2)

# коеффициентуудыг олвол,
a1 <- (Mxy - Mx*My)/(Mxx - Mx**2)
a2 <- n*(Mx**2*My - Mx*Mxy)/(Mxx - Mx**2)

# регрессын тэгшитгэлийг бодож зурвал,
y_lin  <- a1*x + a2
plot(x,y,type="l",col=2)
lines(x,y_lin,col=3)
```

*** =sct
#no SCT

--- type:NormalExercise lang:r xp:150 skills:1 key:da3ca880f5
## Бодлого 19

$N(a,1)$ - нормаль тархалттай $\xi$ - с.х -ний $n=10$ ширхэг элементтэй түүврийн хувьд,

![equation](https://s3.amazonaws.com/smt201-16/render19_1.png) 

**(1)** нь $N(0,1)$ - стандарт нормаль тархалттай бол  **а** - параметрын $p=0.95$ түвшний итгэх завсрыг ол.


*** =instructions

$x, y$ -ын дундаж:

- `Mx <- mean(x)`
- `My <- mean(y)`

$xy, x^2$ -ын дундаж:

- `M(xy) <- mean(x*y)`
- `M(xx) <- mean(x**2)`

Коеффициентуудыг олвол:

- $a1 = \frac{M(xy)-MxMy}{Mx^2-(Mx)^2}$
- $a2 = n \frac{Mx^2My - Mx M(xy)}{Mx^2-(Mx)^2}$
- $ y = a1x + a2$ гэж олдоно.

*** =hint
# no hint

*** =pre_exercise_code
```{r}
x <- c(-1, -0.75, -0.5, -0.25, 0.0, 0.25, 0.5, 0.75, 1.0)
y <- c(2.08, 1.83, 1.57, 1.13, 0.89, 0.75, 0.30, 0.06, -0.01)
n <- 9
```

*** =sample_code
```{r}
# x, y -ын дундаж
Mx <- mean(x)
My <- mean(y)
# xy, x^2 -ын дундаж
Mxy <- mean(x*y)
Mxx <- mean(x**2)

# коеффициентуудыг олвол,
a1 <- (Mxy - Mx*My)/(Mxx - Mx**2)
a2 <- n*(Mx**2*My - Mx*Mxy)/(Mxx - Mx**2)

# регрессын тэгшитгэлийг бодож зурвал,
y_lin  <- a1*x + a2
plot(x,y,type="l",col=2)
lines(x,y_lin,col=3)
```


*** =solution

```{r}
# x, y -ын дундаж
Mx <- mean(x)
My <- mean(y)
# xy, x*x -ын дундаж
Mxy <- mean(x*y)
Mxx <- mean(x**2)

# коеффициентуудыг олвол,
a1 <- (Mxy - Mx*My)/(Mxx - Mx**2)
a2 <- n*(Mx**2*My - Mx*Mxy)/(Mxx - Mx**2)

# регрессын тэгшитгэлийг бодож зурвал,
y_lin  <- a1*x + a2
plot(x,y,type="l",col=2)
lines(x,y_lin,col=3)
```

*** =sct
#no SCT

--- type:NormalExercise lang:r xp:200 skills:1  key:39be7e1e33
## Бодлого 20

**Бодлого 17** -гийн өгөгдөл болон хамгийн бага квадратын аргаар олсон параметруудыг ашиглан $( yk - y_{lin})$ хэмжигдэхүүн $N(0,\sqrt{0.1})$ тархалттай болохыг **Уилкоксоны шинжүүрээр** шалга! Энд стандарт нормаль тархалттай загвар түүврээр **t=(t1, ..., t30)** - г хэрэглэнэ үү. 

*** =instructions

$x, y$ -ын дундаж:

- `Mx <- mean(x)`
- `My <- mean(y)`

$xy, x^2$ -ын дундаж:

- `M(xy) <- mean(x*y)`
- `M(xx) <- mean(x**2)`

Коеффициентуудыг олвол:

- $a1 = \frac{M(xy)-MxMy}{Mx^2-(Mx)^2}$
- $a2 = n \frac{Mx^2My - Mx M(xy)}{Mx^2-(Mx)^2}$
- $ y = a1x + a2$ гэж олдоно.

*** =hint
# no hint

*** =pre_exercise_code
```{r}
# Загвар түүвэр
t <- c(-3240, 1960, -2902, -2512, -2204, -4007, -537, 928, -2026, 5925, 5516, 1905, -4476, 1642, -3848, 2096, 3471, 2028, -5266, -1358, 68, 4554, -389, -160, -51, -1243, 1430, -1772, -414, -426)/1000

# x,y санамсаргүй хэмжигдэхүүний хослол түүвэр
x <- c(-1, -0.75, -0.5, -0.25, 0.0, 0.25, 0.5, 0.75, 1.0)
y <- c(2.08, 1.83, 1.57, 1.13, 0.89, 0.75, 0.30, 0.06, -0.01)

# түүврийн тоо
n <- 9

# x, y -ын дундаж
Mx <- mean(x)
My <- mean(y)

# xy, xx -ын дундаж
Mxy <- mean(x*y)
Mxx <- mean(x**2)

# коеффициентуудыг олвол,
a1 <- (Mxy - Mx*My)/(Mxx - Mx**2)
a2 <- n*(Mx**2*My - Mx*Mxy)/(Mxx - Mx**2)
rm(Mx,My,Mxx,Mxy)
# регрессын тэгшитгэлийг бодох,
y_lin  <- a1*x + a2
#plot(x,y,type="l",col=2)
```

*** =sample_code
```{r}
# x, y -ын дундаж
Mx <- mean(x)
My <- mean(y)
# xy, x^2 -ын дундаж
Mxy <- mean(x*y)
Mxx <- mean(x**2)

# коеффициентуудыг олвол,
a1 <- (Mxy - Mx*My)/(Mxx - Mx**2)
a2 <- n*(Mx**2*My - Mx*Mxy)/(Mxx - Mx**2)

# регрессын тэгшитгэлийг бодож зурвал,
y_lin  <- a1*x + a2
plot(x,y,type="l",col=2)
lines(x,y_lin,col=3)
```


*** =solution

```{r}
# x, y -ын дундаж
Mx <- mean(x)
My <- mean(y)
# xy, x*x -ын дундаж
Mxy <- mean(x*y)
Mxx <- mean(x**2)

# коеффициентуудыг олвол,
a1 <- (Mxy - Mx*My)/(Mxx - Mx**2)
a2 <- n*(Mx**2*My - Mx*Mxy)/(Mxx - Mx**2)

# регрессын тэгшитгэлийг бодож зурвал,
y_lin  <- a1*x + a2
plot(x,y,type="l",col=2)
lines(x,y_lin,col=3)
```

*** =sct
#no SCT
